# Week 0: Exercise introduction and setup <span style="font-size: 0.3em;">v20260825a</span>

This exercise is an *optional* pre-exercise, that will show you how the weekly exercises and assignments are done. 

You can use the exercise to check that you have the sufficient coding skills, as well as ensuring that you have your integrated development environment (IDE) setup.

**Content:**

- Part 0: Downloading Python, Anaconda and your IDE
- Part 1: Getting started with Jupyter Notebooks
- Part 2: Getting started with Python
- Part 3: Matrix operations
- Part 4: Basic plotting
- Part 5: An introduction to Pandas
- Part 6: Recommended materials to catch up on Python

**Objectives:**
- Have setup your Python IDE.
- Understand Jupyter Notebooks and the various shortcuts.
- Write and execute basic Python code.
- Understand how data can be represented as vectors and matrices in numerical Python (NumPy).
- Understand basic ways of plotting data. 
- Understand the basics of Pandas.
- Understand the format of the exercises.

## Introduction:
The exercises in the course are structured such that you work your way through an exercise notebook like this one. Your task will be to fill in missing pieces of code and explanations to make the Notebooks run - you do so by deleting `raise NotImplementedError()` and inserting your own code. 

If you encounter any problems in doing this exercise - in e.g. setting up your IDE - make sure to get in touch with a teaching assistant after the first lecture at the exercises. The first real course exercise (Week 1) will be done after the first lecture. Note that if you have difficulties following the scripts, catching up on your progamming-skills are to be done as a **self-study exercise** (see optional material in the end of this notebook).

We recommend that you create a folder named `02452exercises` dedicated to the course exercises. Each week, you will receive a zip file via DTU Learn containing the associated files, which should be placed within this directory, for example, `02452exercises/week3`.
<br>

### Weekly Assignments:

At the end of each weekly notebook, you will find a weekly **Assignment**-section, consisting of a couple of assignment-tasks. These assignments are part of the course **hand-in requirements**, and you are expected to hand in these weekly notebooks. 

Only the **assignment**-part of the notebooks will be graded, while the **weekly exercises** are there to guide you through the course material, and help you solve the **assignments**.

Some exercises throughout the notebooks will be marked as *optional*. These are not needed for successfully solving the weekly assignment, so you should prioritize accordingly and are welcome to skip exercises if needed.

---

## Part 0: Downloading Python, Anaconda and your IDE

There are many IDEs on the market that vary from being Python-specific to general multilanguage IDEs. If you already have your favourite IDE setup, you are very welcome to keep that. In this course, we recommend using [Visual Studio Code](https://code.visualstudio.com/).

[**DTU Python support**](https://pythonsupport.dtu.dk/install/python.html) has provided an easy guide on getting started, including an approach of downloading Python, Anaconda and Visual Studio in a single step.

**Task 0.1:** If you do not already have Python, Anaconda and an IDE - click the aforementioned link to DTU Python support, and follow the guide for your Operating System.

<br>

---

## Part 1: Getting started with Jupyter Notebooks

**Creating an environment with Anaconda**

We strongly recommend setting up a Python virtual environment to encapsulate the packages you need for this course - this will prevent them from interfering with other Python installations and environments. 
We recommend setting up an environment using [Anaconda](http://www.anaconda.com/download/) - which you should already have installed if you followed Part 0.

**When working within an activated virtual environment, whatever packages you install are only within this environment and do not conflict with the Python version of the system.**

<br>

**Task 1.1:** Create a course-specific virtual environment. 

> *Hint*: Copy the following into a terminal: `conda create --name dtu02452 python=3.11`

> *Hint*: Here "dtu02452" is the name of your environment (you can change this if you want). Remember that you should use a course-compatible Python version 3.8-3.11 (we recommend 3.11).

Your terminal should show something like
`(base) YOU_NAME@MacBookPro ~ %`,

<br>

**Task 1.2:** Try activating your new environment.

> *Hint:* Copy the following into a terminal: `conda activate dtu02452`

Your terminal should now have changed to something like
`(dtu02452) YOU_NAME@MacBookPro ~ %`,
which means that the environment `dtu02452` is activated. If you install a Python package it is now installed within the environment and does not conflict with the Python of your base system.

You are able to install packages within your activated environment by running the command:

`pip install NAME_OF_PACKAGE` or `conda install NAME_OF_PACKAGE`

<br>

**Task 1.3:** Install the required course-specific packages, by pip installing from the requirements file. 

> *Hint:* You can install a collection of packages at once. Copy the following into a terminal: `pip install -r requirements.txt`

> *Hint:* Make sure that the `requirements.txt`-file is in your current directory.

To deactivate the environment run in terminal: `conda deactivate`.

<br>

**Setting up Visual Studio Code IDE**

Visual Studio Code is a powerful and versatile code editor. It supports many programming languages and offers helpful features such as syntax highlighting, auto-completion, debugging, and extensions to add new functionality - to get started with Notebooks in VSCode, we need to install two extensions. 

**Task 1.4:** Open Visual Studio Code, and install the extensions for Python and Jupyter.

> *Hint:* You can search for extensions in the extensions tab on the left. 

<br>

Your VSCode installation should now be ready for handling Notebooks.

**Task 1.5:** Open the `.ipynb`-version of this Notebook in VSCode, click on `Select Kernel` in the upper-right corner, and select the Virtual Environment you previously created i.e., `dtu02452`. If you have successfully created the environment, when selecting the kernel, you should see something like: 

`dtu02452 (Python 3.11.13) ~/opt/anaconda3/envs/dtu02452/bin/python`.

**Task 1.6:** Try running the cell below, to check whether your installation has worked.
> *Hint:* Press `Shift + Enter` or click the `▷`-icon.

```python
print("Hello World!")
```

**Task 1.7:** Combine text and code in a print-statement. 
> *Hint:* You can print both words and math in the same line using `print(f"...")`. <br> Anything inside `{ }` will be calculated before it gets printed. 

> *Hint:* `\n` inserts a new line.

```python
print(f"2 + 2 \n= {2 + 2}")
```

As you might have noticed, Jupyter Notebooks is a tool that combines Python code and Markdown text into a single document, by creating different types of cells. Code cells let you write and run Python, while Markdown cells are used for formatted text, explanations, and equations. 

**Task 1.8:** Create one code cell and one markdown cell using keyboard shortcuts.

> *Hint:* Press `b` to create a cell below, `m` to turn it into a markdown cell, and `y` to turn it back into a code cell.

> *Hint:* Press `x` to delete a cell.

You can navigate cells using the arrow keys, and edit cells by pressing `enter`. 

<br>


---

## Part 2: Getting started with Python
In Python you need to 'import' packages and external functions before you can use them. We can import NumPy (which enables us to work with matrices, among other things) by writing `import numpy as np`.

We load the package into the "namespace" `np` to reference it easily, now we can write `np.sum(X)` instead of `numpy.sum(X)`.

**Task 2.1:** (*Making sequences*) We often need to use a sequence of numbers. Observe the various results obtained when running the following code cell:

```python
# Loading numpy
import numpy as np

# define variable a with numbers in the range from 0 to 7 (not inclusive)
a = np.arange(start=0, stop=7)

# define variable b with numbers in the range from 2 to 17 in steps of 4
b = np.arange(start=2, stop=17, step=4)

# similar to b but without explicit decleration of the input arguments names
c = np.arange(100, 95, -1) 
d = np.arange(1.2, 1.9, 0.1)
e = np.pi * np.arange(0, 2.5, 0.5)

# print the variables
print(f"a: {a}")
print(f"b: {b}")
print(f"c: {c}")
print(f"d: {d}")
print(f"e: {e}")
```

**Task 2.2:** (*Indexing*) We often need to retrieve or assign a value to a certain part of a vector or matrix. Inspect the cell below to see how to do indexing in Python.

> *Hint:* Python uses zero-indexing.

> *Hint:* You can run `help(insert_function_name_here)` to get information on what a function does. 


```python
# Extracting the elements from vectors is easy. Consider the following definition of x and the results
x = np.concatenate([np.zeros(2), np.arange(0, 3.6, 0.6), np.ones(3)])
print(f"x: {x}")  # Print the vector x
print(f"x[1:5]: {x[1:5]}")  # take out elements 2 through 5 (notice 6 is not included)
print(f"np.size(x): {np.size(x)}")  # Print the size of x (equivalent to len(x) since x is an array)
print(f"len(x): {len(x)}")  # Print the length of x

print(f"x[-1]: {x[-1]}")  # take the last element of x
print(f"x[1::2]: {x[1::2]}")  # return every other element of x starting from the 2nd (due to zero-indexing)

```

**Task 2.3:** The length of `x` is 11. Create a cell below and run `x[11]`. What happens - and why?
- *Answer:*

Inserting numbers into vectors is also easy. Notice that we're inserting the same scalar value "pi" into all elements that we index `x` with.

```python
x[1::2] = np.pi # set every other element of x starting from the 2nd to pi
print(f"x after inserting pi:\n{x}")
```

---

## Part 3: Vector and Matrix operations

**Task 3.1:** We will use various vector/matrix operations extensively throughout the course. <br>
Inspect the cell below to see how to create various arrays in Python.

```python
# Setup three arrays
x = np.arange(1, 6)
y = np.arange(start = 2, stop = 12, step = 2)
z = np.array([1, 2, 3, 4])

print(f"x: {x}")
print(f"y: {y}")
print(f"z: {z}")
```

Arrays and matrices are two data structures added by NumPy package to the list of basic data structures in Python (lists, tuples, sets). We shall use both array and matrix structures extensively throughout this course, therefore make sure that you understand differences between them (multiplication, dimensionality).

Generally speaking, array objects are used to represent scientific, numerical, N-dimensional data. While matrix objects can be very handy when it comes to algebraic operations on 2-dimensional matrices.

**Task 3.2:** Look at the cell below to see how to create matrices explicitly and reshape an array into a matrix.

```python
A1 = np.array([[1, 2, 3], [4, 5, 6]])  # define explicitly
A2 = np.arange(1, 7).reshape(2, 3)  # reshape a list of numbers into a 2x3 matrix
```

You will likely be doing a lot of matrix and vector transposes - there are many ways of doing these. 

**Task 3.3:** Inspect the cell below to see various ways of transposing a numpy object.
> *Hint:* `y.T` is the most common way to transpose, though they are all equal.

```python
print("Vector transpositions:\n")
print(f"np.transpose(y):\n{np.transpose(y)}")  # transposition/transpose of y
print(f"y.transpose():\n{y.transpose()}")      # also transpose
print(f"y.T:\n{y.T}")                          # also transpose

print("\nMatrix transpositions:\n")
print(f"np.transpose(A1):\n{np.transpose(A1)}")  # transposition/transpose of A1
print(f"A1.transpose():\n{A1.transpose()}")      # also transpose
print(f"A1.T:\n{A1.T}")                          # also transpose
```

You will also be doing a lot of matrix-vector multiplications.

**Task 3.4:** Inspect the code below to see how these are done:

```python
print(f"np.multiply(x,y): {np.multiply(x, y)}")  # element-wise multiplication
print(f"np.dot(x, y.T): {np.dot(x, y.T)}")       # matrix multiplication
print(f"x @ y.T: {x @ y.T}")                     # also matrix multiplication
```

**Task 3.5:** As you might know, matrices come in many forms - inspect common matrices below:

```python
A1 = np.array([[1, 2, 3], [4, 5, 6]])  # define explicitly
A2 = np.arange(1, 7).reshape(2, 3)  # reshape a list of numbers into a 2x3 matrix
A3 = np.zeros([3, 3])  # zeros array of size 3x3
A4 = np.eye(3)  # diagonal array of size 3x3
A5 = np.random.rand(2, 3)  # random array of size 2x3 (range 0-1)
A6 = A1.copy()  # copy
A7 = A1  # alias

print(f"A1:\n{A1}\n")
print(f"A2:\n{A2}\n")
print(f"A3:\n{A3}\n")
print(f"A4:\n{A4}\n")
print(f"A5:\n{A5}\n")
print(f"A6:\n{A6}\n")
print(f"A7:\n{A7}")
```

**Task 3.6:** It is easy to extract and/or modify selected items from arrays/matrices. Inspect how to index matrix elements:

```python
M = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])  # define explicitly

print(f"first element, M[0,0]: \n{M[0, 0]}\n")  # first element
print(f"last element, M[-1,-1]: \n{M[-1, -1]}\n")  # last element
print(f"first row, M[0,:]: \n{M[0, :]}\n")  # first row
print(f"second column, M[:,1]:\n{M[:, 1]}\n")  # second column
print(f"view on selected rows&columns, M[1:3,-1]:\n{M[1:3, -1]}\n")  # view on selected rows&columns
```

Similarly, you can selectively assign values to matrix elements or columns.

**Task 3.7:** Inspect how to assign values in a matrix.

```python
print(f"M before selective assignment:\n{M}\n")


# Similarly, you can selectively assign values to matrix elements or columns:
M[-1, -1] = 10000 # Setting last element to 10000
M[:, 0] = 0 # Setting first column to 0
M[0, 1:] = 1 # Setting first row, except first column to 1

print(f"M after selective assignment:\n{M}")
```

Sometimes we want to index or change elements of a matrix, that fulfill a certain constraint. In these cases we can do *Logical indexing*.

**Task 3.8:** Inspect how we set all elements below 5 to 0.

```python
print(f"M[M > 0.5]:\n{ M[M > 0.5] }\n")  # display values in M that are larger than 0.5
print(f"M[M < 5]:\n{ M[M < 5] }\n")  # display values in M that are less than 5

M[M < 5] = 0  # set all elements that are less than 5 to 0
print(f"M after setting elements < 5 to 0:\n{M}")
```

Below, several examples of common matrix operations, most of which we will use in the following weeks.

**Task 3.9:** First, see how we define two matrices:

```python
M1 = 10 * np.ones([3, 3]) # 3x3 matrix with all elements equal to 10
M2 = np.random.rand(3, 3) # 3x3 matrix with random elements

print(f"M1:\n{M1}\n")
print(f"M2:\n{M2}")
```

**Task 3.10:** See how we do matrix summation, products and comparisons:

```python
print(f"M1 + M2:\n{M1 + M2}\n")  # matrix summation
print(f"M1 @ M2:\n{M1 @ M2}\n")  # matrix product

# Two ways of doing element wise multiplication (Hadamard product)
print(f"M1 * M2:\n{M1 * M2}\n")  # element-wise multiplication
print(f"np.multiply(M1, M2):\n{np.multiply(M1, M2)}\n")  # element-wise multiplication
print(f"M1 > M2:\n{M1 > M2}")  # element-wise comparison
```

We might want to combine / concatenate matrices horizontally and vertically.

**Task 3.11:** Inspect how `np.hstack` and `np.vstack` are used.

```python
M3 = np.hstack((M1, M2))  # combine/concatenate matrices horizontally
# note that this is not equivalent to e.g.
#   l = [M1, M2]
# in which case l is a list, and l[0] is M1
M4 = np.vstack((M1, M2))  # combine/concatenate matrices vertically
print(f"M3:\n{M3}\n")
print(f"M4:\n{M4}")
```

Various useful functions for matrices:

```python
print(f"M3.shape: {M3.shape}\n")  # shape of matrix
print(f"M3.mean(): {M3.mean()}\n")  # mean value of all the elements
print(f"M3.mean(axis=0):\n{M3.mean(axis=0)}\n")  # mean values of the columns
print(f"M3.mean(axis=1):\n{M3.mean(axis=1)}\n")  # mean values of the rows
print(f"np.linalg.inv(M2):\n{np.linalg.inv(M2)}")  # compute inverse matrix
```

---

## Part 4: Basic plotting

It is important to visualize the results you obtain. In this part of the exercise, we will make two plots, a simple, and a more elaborate example. Going forwards, try to keep the elements of the figure made in Task 4.3 in mind as a model for minimum required considerations when making a figure.

**Task 4.1:** Inspect the code cell below to see how to do basic plotting in Python.

```python
import matplotlib.pyplot as plt
import numpy as np
# Plotting style
import seaborn as sns
sns.set_style('darkgrid')
sns.set_theme(font_scale=1.)

# Array of x values
x = np.arange(start = 0, stop = 1, step = 0.1)
# Compute f(x) = exp(x)
f = np.exp(x)

# Plotting
plt.figure(1)
plt.plot(x, f)
plt.xlabel("x") # x-axis label
plt.ylabel("f(x)=exp(x)") # y-axis label
plt.title("The exponential function") # title
plt.show()

```

Now, let us consider a slightly more advanced plot.

We simulate measurements from two sensors (sensor 1 and sensor 2) for a period of 10 seconds, and we say that the sensors output measurents in millivolt.  Since the two measurements are done at the same time, we want to plot them in the same figure. Furthermore, we want to make sure that the axes are labelled correctly both with a name and a designation of the unit of measurement. We also need to make sure that it is easy to see which curves comes from which sensor. Lastly, we ensure that the axes are readable (large enough), both when looking at them in the IDE, but also in an exported version that we might use in a report about the sensors.

**Task 4.3:** Inspect the cell below.

```python
# We simulate measurements every 100 ms for a period of 10 seconds
t = np.arange(0, 10, 0.1)

# The data from the sensors are generated as either a sine or a cosine
# with some Gaussian noise added.
sensor1 = 3 * np.sin(t) + 0.5 * np.random.normal(size=len(t)) # Draw len(t) samples from a normal distribution
sensor2 = 3 * np.cos(t) + 0.5 * np.random.normal(size=len(t)) # Draw len(t) samples from a normal distribution

# Change the font size to make axis and title readable
font_size = 15
plt.rcParams.update({"font.size": font_size})


# Start plotting the simulated measurements
plt.figure(1)
# Plot the sensor 1 output as a function of time, and
# make the curve red ("r") and fully drawn ("-")
plt.plot(t, sensor1, "r-")

# Plot the sensor 2 output as a function of time, and
# make the curve blue ("b") and dashed ("--")
plt.plot(t, sensor2, "b--")

# Ensure that the limits on the axis fit the data
plt.axis("tight")

# Add a legend describing each curve, place it at the "best" location so as to minimize the amount of curve it covers
# Define the name of the curves
legend_strings = ["Sensor 1", "Sensor 2"]
plt.legend(legend_strings, loc="best")

# Add labels to the axes
plt.xlabel("Time [s]")
plt.ylabel("Voltage [mV]")

# Add a title to the plot
plt.title("Sensor outputs")

# Optionally save / export the figure
# plt.savefig("task4_3.png")

# Show the figure in the notebook
plt.show()
```

---

## Part 5: An introduction to Pandas

Pandas is a Python library for working with structured data. Pandas works through DataFrames, a table-like data structure with labeled rows and columns, which makes it easy to:
- Load data from files (CSV, Excel, etc.)
- Inspect and clean datasets
- Select, filter, and transform data
- Compute statistics and summaries quickly

Pandas is widely used in data science and machine learning because it lets you handle datasets efficiently while being easy to visualize.

We will start by importing Pandas and loading the [Titanic dataset](https://www.kaggle.com/competitions/titanic).
This part of the exercises has been borrowed from [Raj Mehrotra](https://www.kaggle.com/code/rajmehra03/a-complete-pandas-tutorial).

```python
# Importing pandas
import pandas as pd

# Load data from csv (Loading data will be explained further in Week1)
df = pd.read_csv(r'titanic.csv')
```

### 5.1: The basics

Pandas lets us get a quick overview of our data, by easily being able to displaying the first and last 5 entries in our dataset through `df.head()` and `df.tail()`.

```python
# See the first 5 rows
df.head()
```

```python
# last 5 rows.
df.tail()
```

We can access the size of the dataset through `df.shape`

```python
# n_samples x n_features
df.shape
```

```python
#List of all the columns
df.columns
```

```python
# Rows index
df.index
```

```python
# Values with their counts in a particular column
df['Pclass'].value_counts()
```

Pandas also let's us quickly get summary statistics of our dataset.

```python
# General description of dataset.
df.describe()
```

### 5.2: Creating DataFrames

We can create empty DataFrames that lets us insert values after the fact.

```python
# empty data frame
df_empty=pd.DataFrame()
df_empty.head()
```

We can also create DataFrames from Python Dictionaries - a very valuable data structure.

If you'd like to know more about Python dictionaries, you can read up on them [here](https://www.w3schools.com/python/python_dictionaries.asp).

```python
student_dict={'Name':['A','B','C'],'Age':[24,18,17],'Roll':[1,2,3]}
df_student=pd.DataFrame(student_dict).reset_index(drop=True) # without this adds an additional index column in df
df_student.head()
```

### 5.3: Treating null values

Despite what common datasets might let you believe, data in the real world is oftentimes messy, containing missing values, outliers and other things making the data problematic to work with. 

We can explore the amount of null values in our dataset.

```python
df.isnull().sum()
```

Or do it directly on a column. 

```python
df['Age'].isnull().sum()
```

A common practice for treating Null values is to impute them with the mean-value of the specific feature - other practices include Median- and Mode-imputation or simply deleting entries with Nulls.

```python
# Impute with the mean value
df['Age'] = df['Age'].fillna(df['Age'].mean())
df['Age'].isnull().sum()
```

```python
# 'Sex' is a categorical value, we can impute with the mode (The most frequently occurring value)
df['Sex'] = df['Sex'].fillna(df['Sex'].mode())
df['Sex'].isnull().sum()
```

### 5.4: Modify/Add new column(s)

In machine learning, we often convert string values into numerical categories as most algorithms only work with numbers, not text labels.

The cell below modifies the existing column, by mapping the string values into binary. 

```python
df['Sex'] = df['Sex'].map({"male":'0',"female":"1"})
df.head()
```

We can also create new columns based on existing columns. The code below extracts the first and last names from the `Name`-column. 

```python
df['last_name'] = df['Name'].apply(lambda x: x.split(',')[0]) # Split by comma and take the first part
df['first_name'] = df['Name'].apply(lambda x: ' '.join(x.split(',')[1:])) # Join back the remaining parts as first name

df.head()
```

Pandas lets us run functions on all entries of a column through `.apply()`. This is preferred over iterating through the items, as it is much faster. 

```python
def findAgeGroup(age):
    if age<18:
        return 1
    elif age>=18 and age<40:
        return 2
    elif age>=40 and age<60:
        return 3
    else:
        return 4
    
# Calling a custom function.
df['Age_group'] = df['Age'].apply(lambda x: findAgeGroup(x))

df.head()
```

### 5.5: Deleting and renaming columns

You will be deleting ("dropping") columns from your datasets a lot throughout the course, especially when splitting a DataFrame into Features and Target variables. 

```python
df = df.drop(['PassengerId'],axis=1) # Removing PassengerId column
df.head()
```

Renaming columns is once again done through a Dictionary.

```python
# Lets try to rename some columns. 
df=df.rename(columns={'Sex':'Gender','Name':'Full Name','last_name':'Surname','first_name':'Name'})
df.head()
```

### 5.6: Filtering and slicing DataFrames
We will oftentimes want to filter our dataset, the cell below slices the DataFrame such that we only look at entries with `Pclass == 3`.

```python
# All rows with pclass==3
df_third_class = df[df['Pclass'] == 3].reset_index(drop=True)  # w/0 drop=True it actually adds a index column rather.
df_third_class.head()
```

We can also filter on multiple columns, the following filters for women over the age of 60.

Note that all women over 60 survived the Titanic. 

```python
df_aged = df[(df['Age'] > 60) & (df['Gender'] == "1")]
df_aged.head()
```

We can also slice our DataFrame based on specific columns.

```python
# Selecting some columns.
df1=df[['Age','Pclass','Gender']]
df1.head()
```

Or slice the columns for numerical entries only. 

```python
# Select numerical columns only
numerics = ['int16', 'int32', 'int64', 'float16', 'float32', 'float64']

df_num = df.select_dtypes(include=numerics)
df_num.head()
```

Or Categorical columns.

```python
# categorical columns
df_cat=df.select_dtypes(include=['object'])
df_cat.head()
```

You've previously learned how to slice matrices based on indices. 

Thankfully, this is also available in Pandas through `.iloc` (position-based indexing) and `.loc` (label-based indexing). 

Slicing for the first 100 rows & all columns:

```python
# First 100 rows & all columns
df_sub1 = df.iloc[0:100, :]
df_sub1.head()
```

```python
df_sub2 = df.iloc[:250, [1, 8]]  # First 250 rows and only columns 1 and 8
df_sub2.head()
```

`.loc` lets us index based on labels / column-names.

```python
# Gender and age, where age > 50
df_sub4 = df.loc[(df['Age'] > 50), ['Gender', 'Age']]
df_sub4.head()
```

### 5.7: Adding and dropping rows

As data-scientists, we are always on the lookout for more data. We can append more data with `.concat()`.

```python
row = {'Age': 24, 'Full Name': 'Peter', 'Survived': 'Y'}
df = pd.concat([df, pd.DataFrame([row])], ignore_index=True)
# assumes NaN for absent keys(columns)
df.tail()
```

We have previously seen how to drop columns, this also works for rows. 

```python
df=df.drop(df.index[-1],axis=0) # Deletes last row
df.head()
```

### 5.8: Sorting

When exploring our data, it is oftentimes useful to sort it - we can sort by a given column:

```python
# sorting by age say in decreasing order.
df = df.sort_values(by=['Age'], ascending=False) # can specify multiple columns in a list as well.
df.head()
```

### 5.9: Grouping

You can use `df.groupby()` when you need to split your data into categories (by one or more keys) and compute aggregates or transformations per group, which lets you quickly summarize patterns and compare segments.

We can group our observations by gender, and then calculate the survival-rate for each gender. 

> **Hint:** reset_index() turns the group labels back into regular columns for a tidy DataFrame.

```python
gender_groups = df.groupby(by=['Gender'])
gender_groups['Survived'].mean().reset_index()
```

We see that the survival-rate was much higher for women, than for men. Can you think of why?

- *Answer:*

Similarly, we can also group our dataframe by the age groups we created earlier, and easily plot the survival rates per age group. 

```python
age_groups = df.groupby(by=['Age_group'])

survival_chance_per_age = age_groups['Survived'].mean().reset_index()

survival_chance_per_age.plot(x='Age_group', y='Survived', kind='bar')
```

---

## Part 6: Recommended materials to catch up on Python

The following online materials are recommended:

- [Introduction into python environment, syntax and data structures](docs.python.org/tutorial). Recommended reading - sections 1, 2, 3, 4 and 5.
- [Tutorial introducing the scientific computing in Python, array and matrix operations, indexing and slicing matrices](https://numpy.org/doc/stable/).
- [Useful reference to scientific computing in Python if you have previous experience with Matlab programming](docs.scipy.org/doc/numpy-1.15.0/user/numpy-for-matlab-users.html).
- [Documentation and examples related to the matplotlib module](matplotlib.sourceforge.net) - which we shall use extensively throughout the course to visualize data and results.
- [Overview of VSCode IDE for Python code editing/debugging](https://code.visualstudio.com/docs/python/python-tutorial).
- [Series of video tutorials covering basics of Python programming (using a different IDE)](https://www.youtube.com/playlist?list=PL6gx4Cwl9DGAcbMi1sH6oAMk4JHw91mC_).
- [DTUs Introduction to Programming (using Python) for absolute beginners](https://02002.compute.dtu.dk/).

The Python tutorial (sections 1-5) and NumPy tutorial are especially recommended. The more you get acquainted with Python now, the easier it will be for you to solve machine learning problems in the following weeks. You will benefit from this course most if you try to implement the solutions on your own, before checking the correct answers. Nevertheless, if you run into problems, the guidelines and the correct scripts will be always provided for your reference.

-----

# Assignment 0:

As mentioned earlier, each week you will be asked to hand in a weekly assignment. This week’s assignment is optional and will not be graded, but it is included so you can practice the Weekly Assignment procedure.

In some exercises, you will see visible tests that immediately show whether your solution is correct. In other exercises, the tests will be hidden or manually graded, meaning you will not get instant feedback.

The expected “answer” can vary: sometimes it is a variable, sometimes it will be a plot, and other times it will be a written interpretation of a plot or another result.

When you see a code cell marked with `### YOUR CODE HERE` and a `raise NotImplementedError()`, remove the placeholder line and insert your solution code.

We will start importing a few tools and printing out some infomation about your platform for debugging in case there are issues with the notebook or tests below. 

<!-- grade_id: cell-8bbaadac2dda1122, points: 0, grade: true, solution: false -->
```python
import utils as utils, re
utils.reset_marks()
a,b,c,d = utils.platform_info()
print(f"Platform info: <<<{a}:{b}:{c}:{d}>>>")
with utils.marks(0): # If we get to this point, we will show a checkmark to indicate that everything looks fine so far
    assert(True)
```

<br>

**Assignment 0.0:** Fill in your student ID and your full name in the variables below (and ensure that the cell runs):

<!-- grade_id: cell-1cae00ca51b85469, points: 0, grade: false, solution: true -->
```python
student_id = "sXXXXXX" # e.g. s123456
student_typewritten_signature = "Firstname Lastname" # e.g. "Hans Hansen"

# YOUR CODE HERE
raise NotImplementedError()
```

Our first small test is to check that the format of the student id is correct.

<!-- grade_id: cell-806456ca23063825, points: 0, grade: true, solution: false -->
```python
# We include such simple tests (without marks) to make sure you are on the right track
with utils.marks(0): 
    assert(re.compile(r"^s\d{6}$").match(student_id))
    assert(len(student_typewritten_signature)>0)
```

___

**Assignment 0.1:** Subtract matrix A from matrix B, save the result as C 

<!-- grade_id: cell-0d5f6d532fa0c5c4, points: 0, grade: false, solution: true -->
```python
import numpy as np

# Defining matrices:
A = np.array([[1,2,3],[4,5,6]])
B = np.array([[7,8,9],[10,11,12]])

# Assignment 0.1: Subtract matrix A from Matrix B element-wise, save the result as C.
# YOUR CODE HERE
raise NotImplementedError()

print(C)
```

<!-- grade_id: cell-2f805d14ccd67926, points: 1, grade: true, solution: false -->
```python
# Visible test for matrix subtraction
with utils.marks(1):
    assert(utils.check_hash(C, ((2,3), 150.0)))
```

**Assignment 0.2:** Find the shape of matrix C, save it as `C_shape`
> *Hint:* You can use `C.shape()`

<!-- grade_id: cell-295e8c447ab1d5c0, points: 0, grade: false, solution: true -->
```python
# Assignment 0.2: Find the shape of C, save it as C_shape
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-afcb7dda7a382d5a, points: 1, grade: true, solution: false -->
```python
# Hidden test to verify the shape of C (this type of test will be run by assessors after submission)
with utils.marks(1, auto=True, visible=False):            
    print("Hidden test")
```

**Assignment 0.3:** Use Matrix-indexing to change the last element of A to -10.

<!-- grade_id: cell-792540f1e451334b, points: 0, grade: false, solution: true -->
```python
# Assignment 0.3: Set the last element of matrix A to -10
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-379c1c20ff59eb95, points: 2, grade: true, solution: false -->
```python
# Visible test that provides immediate results.
with utils.marks(2):
    assert(utils.check_hash(A, ((2, 3), -1.6527970145028927)))
```

**Assignment 0.4:** Plot a cosine curve for values in the range $[-\pi, \pi]$.

> *Hint:* Take inspiration from **Task 4.1**.

> *Hint:* The value of Pi is available in `np.pi`.

An assignment such as this will be validated based on a visual inspection.

<!-- grade_id: cell-cb19beb145359f36, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt
# Plotting style
import seaborn as sns
sns.set_style('darkgrid')
sns.set_theme(font_scale=1.)

# Assignment 0.4: Plot a cosine curve for values in the range $[-\pi, \pi]$.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b74c6d775fe03fac, points: 2, grade: true, solution: false -->
```python
with utils.marks(2, auto=False, visible=False):
    print("Manual test")
```

**Assignment 0.6:** Calculate the average survival-rate for women and men on the Titanic, based on the pandas dataframe. <br> Save them as `female_survival` and `male_survival`.

> *Hint:* You can slice the dataframe for all women with: `df[(df['Sex'] == "1")]` and all men with: `df[(df['Sex'] == "0")]`

> *Hint:* Whether a person survived is a boolean in the column `Survived`.

> *Hint:* The mean of a list of boolean values, is the fraction of values that are True (i.e. the survival-rate in this case)

<!-- grade_id: cell-4aabdbd1376abcef, points: 0, grade: false, solution: true -->
```python
# Importing pandas
import pandas as pd

# Load data from csv (Loading data will be explained further in Week1)
df = pd.read_csv(r'titanic.csv')
# 'Sex' is a categorical value, we can impute with the mode (The most frequently occurring value)
df['Sex'] = df['Sex'].fillna(df['Sex'].mode())
# We map male to 0 and female to 1
df['Sex'] = df['Sex'].map({"male":'0',"female":"1"})

# YOUR CODE HERE
raise NotImplementedError()

print("Women survival-rate: ", female_survival)
print("Men survival-rate: ", male_survival)
```

<!-- grade_id: cell-e22cc16e0a897843, points: 2, grade: true, solution: false -->
```python
# Visible test
with utils.marks(2):
    assert(utils.check_hash(female_survival, ((), 3.7101910828025475)))
    assert(utils.check_hash(male_survival, ((), 0.9445407279029462)))
```

<br>

**Assignment 0.7:** Why do you think the average survival-rates are so different for the two genders (maximum 100 words)? 
<br>*Note:* These open-ended questions will typically ask you to reflect on key aspects of the problem, the code you wrote, the analysis you performed, or the results you obtained.

*Answer:* [Insert your answer here]

<!-- grade_id: cell-c0fbcf3bf8758a0e, points: 2, grade: true, solution: false -->
```python
with utils.marks(2, auto=False, visible=False):
    print("Manual test")
```

<br>

**Assignment 0.8:** In some weeks, the weekly assignment may include tasks similar to exam questions. 
<br>The goal is to help you get comfortable with the format and the level of reasoning expected at the exam. 

<br>Since this is Week 0, you’re lucky - there are no exam-style tasks this time.

<br>

**Assignment 0.9:** Declaration on the use of AI for solving this assignment (mandatory)

Answer by replacing [your answer here] with your response.
1. *I abide by [DTU's code of honour](https://student.dtu.dk/en/exam/exam-cheating/dtu-code-of-honour) and take full responsibility for the content of this submission (yes / no)*: 
    - [your answer here] 
2. *To what extent did you use generative AI to solve this assignment (0-100%)*:
    - [your answer here] 
3. *What was the primary use of generative AI, if any (writing / coding / other)*: 
    - [your answer here] 
4. *I feel I understand the key techniques/algorithms/methods/coding elements used in this exercise such that I can apply it in a **no aids** exam (yes/no)*: 
    - [your answer here] 

<!-- grade_id: cell-69b38d4c0bdcdc4c, points: 0, grade: true, solution: false -->
```python
with utils.marks(0, auto=False, visible=False):
    print("Manual test")
```

**Assignment 0.10**: Other comments, issues or suggestions (optional)

*Answer (optional)*: [your answer here]

<br>

**Assignment 0.11:** To submit this notebook, make sure you have run everything in the **Assignment**-part, and convert this notebook to an HTML file.

> *Hint:* Open the **command palette** in Visual Studio Code, by pressing `Cmd + Shift + P` on Mac or `Ctrl + Shift + P` on Windows. 

> *Hint:* Search for `Jupyter: Export to HTML` and save the HTML file. 

> *Hint:* If you are running the notebook in the browser via the Jupyter interface, go to `File`, then `Save` and choose `Save and Export Notebook` as and select HTML.

**Assignment 0.12:** Hand in your `.ipynb` and `.HTML` file under assignments on DTU Learn for the relevant week.

**<h4>Summary of the points</h4>**

```python
utils.marks_summary()
```